# Assignment 3

This assignment gives you practice working with dictionaries and functions.

# Part 1: Dictionaries

Dictionaries map keys to values. Like sets, they must have unique keys, and setting a key twice discards any previous value and keeps the latest one. First let's test your understanding of these concepts:

Consider the following dictionary:

```python
d = {'a': 1,
     'b': 2,
     'a': 3}
```

What is the length of `d`?

What is the value mapped by the key `'a'`?

<!-- grade_id: cell-668140db86be239d, points: 2, grade: true, solution: true -->
The length of d is 2. The value mapped by 'a' is 3.

These properties of dictionaries allow it, like sets, to be very fast for certain operations, such as key-value lookup and assignment. Let's demonstrate this by creating our own frequency distribution (like `nltk.FreqDist`). First we will see how long it takes to construct the frequency distribution using list methods. Since we want to pair each unique word with its count, we go over the set of words and use the `count()` method on the original list to see how many times it appears.

```python
import time
import nltk

alice_words = nltk.corpus.gutenberg.words('carroll-alice.txt')
words_to_count = alice_words[:5000]
unique_words = set(words_to_count)

print('Version 1: using list.count() (give it a few seconds)')
start = time.perf_counter()                # get the start time

fd1 = {}                                   # make a dictionary to hold the frequency distribution
for word in unique_words:                  # only loop over unique words so we don't count each more than once
    fd1[word]= words_to_count.count(word)  # set the word to its count in the original text

stop = time.perf_counter()                 # get the end time
print('Version 1 took {} seconds'.format(stop - start))
print('Alice appears {} times in the first 5000 words'.format(fd1['Alice']))
```

Output:
```
Version 1: using list.count() (give it a few seconds)
Version 1 took 4.5339145270008885 seconds
Alice appears 50 times in the first 5000 words
```

Now to compare, create the frequency distribution using dictionary methods. Here you will count the words incrementally by adding one to each words count each time it is encountered. For this to work, the for-loop needs to go over all words, not just the unique words.

**hint:** inside the loop you will assign `fd2[word]` to its previous value plus 1. The first time the word is encountered it will not be in the dictionary so you need to detect this and give it an initial value. How you do this is up to you (e.g., an `if` block or a dictionary method like `get()`, etc.).

<!-- grade_id: cell-bb775b8f526a3ea5, points: 0, grade: false, solution: true -->
```python
print('Version 2: incrementally build a dictionary')

start = time.perf_counter()             # get the start time

fd2 = {}
for word in words_to_count:
    if word in fd2:
        fd2[word] = fd2[word] + 1
    else:
        fd2[word] = 1
    
stop = time.perf_counter()              # get the end time
print('Version 2 took {} seconds'.format(stop - start))
print('Alice appears {} times in the first 5000 words'.format(fd2['Alice']))
```

Output:
```
Version 2: incrementally build a dictionary
Version 2 took 0.018007074999331962 seconds
Alice appears 50 times in the first 5000 words
```

<!-- grade_id: cell-6b0b6e7591330ff4, points: 1, grade: true, solution: false -->
```python
# ensure fd2 is the same as fd1
assert fd1 == fd2

print('All tests passed!')
```

Output:
```
All tests passed!
```

Why is Version 2 so much faster than Version 1?

<!-- grade_id: cell-f9fa8bbceb783caf, points: 1, grade: true, solution: true -->
In version 1, python has to go through the entire list each time in order to determine if the target word is present (and therefore the program has to run multiple times) whereas in version 2, python only has to go through the dictionary once and add 1 to the existing value of fd2 when the target word is located. 

Consider the following situations and say for each whether it is better to use a list or a dict.

1. You are creating a part-of-speech tagged corpus where each word has an associated part-of-speech tag, e.g.:

       The/DT  artist/NN  will/MD  record/VBZ  a/DT  new/JJ  record/NN
   
   You will use these to model syntax with POS-tag bigrams. Do you store these as a list of pairs (`[('The', 'DT'), ...]`) or as a dictionary (`{'The': 'DT', ...}`)?
   
2. You want to create a baseline POS-tagger which is used to assign the most frequent tag to a given word. For example, you want a function that takes a word and returns a tag.

   Do you store these as a list of pairs (`[('record', 'VBZ'), ...]`) or a dictionary (`{'record': 'VBZ', ...}`)?

<!-- grade_id: cell-39ebd59e359361d2, points: 2, grade: true, solution: true -->
1. List (you don't want to lose the word order information)
2. Dict

# Part 2: Functions

Functions are useful so you only have to write frequently used code once and then you can call it from other places in your code later. This helps reduce work and bugs. Write a function named `is_empty()` that takes Python data structures like lists and returns `True` if the structure is empty and `False` if not.

<!-- grade_id: cell-fa6a99386a4668d1, points: 0, grade: false, solution: true -->
```python
def is_empty(s):
    if len(s) == 0:
        return True
    else:
        return False
```

<!-- grade_id: cell-4865a80903af6d63, points: 2, grade: true, solution: false -->
```python
assert is_empty([]) == True
assert is_empty({}) == True
assert is_empty([1,2]) == False
print('All tests passed!')
```

Output:
```
All tests passed!
```

Functions must reach a `return` statement to return a value other than `None`. Once a `return` statement is reached, the function exits immediately. Please fix the bugs (hint: there are 2 bugs) in the following function:

```python
def word_lengths(words):
    """Return a list of word lengths for each word in words."""
    lengths = []
    if len(words) > 0:
        for word in words:
            lengths.append(len(word))
    return lengths
```

<!-- grade_id: cell-ebf36ea9bd6751ff, points: 2, grade: true, solution: false -->
```python
assert word_lengths(['The', 'dog', 'barks', '.']) == [3, 3, 5, 1]
assert word_lengths([]) == []
assert word_lengths(['']) == [0]
print('All tests passed!')
```

Output:
```
All tests passed!
```