Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = "Teo Kai Wen"
```

---

# Assignment 1

There are two coding parts and one short answer part to this assignment.

For these you will use the first verse and refrain from the Mexican folk song [*De Colores*](https://en.wikipedia.org/wiki/De_Colores). Run the following to import the NLTK and make the `de_colores` variable available later:

```python
import nltk

# In Python, a string with 3 quotes can span multiple lines
de_colores = '''
De colores, de colores
Se visten los campos en la primavera.

De colores, de colores
Son los pajaritos que vienen de afuera.

De colores, de colores
Es el arco iris que vemos lucir.

Y por eso los grandes amores
De muchos colores me gustan a mí.

Y por eso los grandes amores
De muchos colores me gustan a mí.
'''
```

## Preface: Making Functions

We have not yet covered how to create Python functions, but consider the following function that adds two numbers:

```python
def add(x, y):
    z = x + y
    return z
```

After `def` is the **function name** (`add`; what you use to call it later). Inside the parentheses are the **arguments** (`x` and `y`), which are variables that are available inside the function. The indented part is the **function body**. At the end of the function body is a **return statement** that specifies the **return value** after `return`.

For simple functions you can put the interesting part on the return statement and avoid creating some new variables. For instance, the following is equivalent:

```python
def add(x, y):
    return x + y
```

## Part 1: Case-normalization

In this part you will explore how case normalization affects frequency distributions by writing functions to return normalized tokens. You may find the following string methods useful:

- [str.split()][] -- break up a string at whitespace and return a list of tokens
- [str.lower()][]/[str.upper()][] -- return a string with only lower-/upper-case letters

[str.split()]: https://docs.python.org/3/library/stdtypes.html#str.split
[str.lower()]: https://docs.python.org/3/library/stdtypes.html#str.lower
[str.upper()]: https://docs.python.org/3/library/stdtypes.html#str.upper

Note: [str](https://docs.python.org/3/library/stdtypes.html#str) is just the string type in Python. When you call these methods, you replace `str` with an actual quoted string or string variable, e.g., `'hi'.upper()` or `my_string.split()`.

For example, run the following to see how they work (this part is not graded):

```python
'one two three'.split()
```

Output:
```
['one', 'two', 'three']
```

```python
print('LoWeR'.lower(), 'UpPeR'.upper())
```

Output:
```
lower UPPER
```

Now wrap these string methods in functions that work on a string argument. I have provided the scaffolding for the functions. All you need to provide is the function body and a return statement. (this part is graded)

<!-- grade_id: cell-c8248f64eeac8e49, points: 0, grade: false, solution: true -->
```python
def tokenize(s):
    """Return the list of all the tokens in string *s*."""
    return s.split()
    raise NotImplementedError()

def normalize(s):
    """Return the case-normalized string for *s*."""
    return s.lower()
    raise NotImplementedError()
```

```python
normalize(de_colores)
```

Output:
```
'\nde colores, de colores\nse visten los campos en la primavera.\n\nde colores, de colores\nson los pajaritos que vienen de afuera.\n\nde colores, de colores\nes el arco iris que vemos lucir.\n\ny por eso los grandes amores\nde muchos colores me gustan a mí.\n\ny por eso los grandes amores\nde muchos colores me gustan a mí.\n'
```

Now run the following to ensure the tests pass:

<!-- grade_id: cell-d8a9c800a71ec7ee, points: 2, grade: true, solution: false -->
```python
assert tokenize('one') == ['one']
assert tokenize('one two') == ['one', 'two']
assert tokenize('two one') == ['two', 'one']
assert tokenize('') == []
assert len(tokenize(de_colores)) == 59
```

<!-- grade_id: cell-d6e2f1c72f972d80, points: 2, grade: true, solution: false -->
```python
assert normalize('ABC') == 'abc'
assert normalize('abc') == 'abc'
assert normalize('') == ''
assert normalize(de_colores).split()[0] == 'de'
```

Now use `nltk.FreqDist` and the functions you wrote above to create a frequency distribution of the raw (not normalized) tokens in a string.

<!-- grade_id: cell-f32c7a112635ce41, points: 0, grade: false, solution: true -->
```python
from nltk.book import *

def token_frequencies(s):
    """Return an nltk.FreqDist of the tokens in string *s*."""
    return FreqDist(tokenize(s))
    raise NotImplementedError()
```

<!-- grade_id: cell-1f63d520cfd07a34, points: 2, grade: true, solution: false -->
```python
raw = token_frequencies(de_colores)
assert raw['De'] == 5
assert raw['de'] == 4
assert raw['colores'] == 5
assert raw['colores,'] == 3
```

Now create a frequency distribution using normalized tokens:

<!-- grade_id: cell-5c59ebdda7530d34, points: 0, grade: false, solution: true -->
```python
def normalized_token_frequencies(s):
    """Return an nltk.FreqDist of the tokens in string *s*."""
    return FreqDist(normalize(s).split())
    raise NotImplementedError()
```

<!-- grade_id: test_normalized_tokens, points: 2, grade: true, solution: false -->
```python
normalized = normalized_token_frequencies(de_colores)
assert normalized['de'] == 9
assert normalized['De'] == 0
assert normalized['colores'] == 5
assert normalized['colores,'] == 3
```

Notice the different frequencies between the raw and normalized distributions. Why does normalizing change the distribution?

<!-- grade_id: cell-41a2a9d6ece210c5, points: 1, grade: true, solution: true -->
Because normalizing collapses capitalized and non-capitalized versions of the same word into one category, unlike raw distributions, which regard the two as distinct.

Look at tokens whose frequencies did not change. What other kinds of normalization might help you to get more accurate word frequencies? (Alternatively: what normalization strategies might help languages other than Spanish or English?)

<!-- grade_id: cell-d4f7bf38a8cda2cc, points: 1, grade: true, solution: true -->
Split the data into individual words by using a function like word_tokenize rather than using spaces to define word boundaries. For other languages, it might be necessary to have access to a something like a corpus that the data can be checked against so that multiple characters that form a single word are not split up, and other features unique to a language (like particles) will not be overlooked.