Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Teo Kai Wen"
COLLABORATORS = ""
```

---

# Assignment 7

This assignment goes a little deeper into the basic concepts covered in previous weeks and deals with advanced functions and testing.

## Q1 (1 point)

Consider:

```python
>>> x = [1,2,3]
>>> y = x
>>> z = [1,2,3]
>>> x == y
True
>>> x == z
True
>>> x is y
True
>>> x is z
False
```

Why does `x is y` evaluate as `True` while `x is z` evaluates as `False`?

<!-- grade_id: cell-d941e8c3aa6d124f, points: 1, grade: true, solution: true -->
'==' checks for value equality whereas 'is' checks whether both operands refer to the same object.

## Q2 (2 points)

Transform the list comprehension below into normal `for` loops. Accumulate the results into a list named `words2`. You can reuse the imported `re` and the `lines` list for your answer.

```python
import re
lines = [
    '¡Que llueva! ¡Que llueva!',
    'El quetzal está en la cueva.'
]
words = [word for line in lines for word in re.findall(r'\w+', line)]
words
```

Output:
```
['Que',
 'llueva',
 'Que',
 'llueva',
 'El',
 'quetzal',
 'está',
 'en',
 'la',
 'cueva']
```

<!-- grade_id: cell-a724165575ada809, points: 0, grade: false, solution: true -->
```python
words2 = []
words2str = ' '.join(lines)
words3str = ''.join(c for c in words2str if c not in ('¡','!','.'))
newword1 = words3str.split()
for newword in words3str:
    words2.extend(newword1)
    print (newword1)
    break
```

Output:
```
['Que', 'llueva', 'Que', 'llueva', 'El', 'quetzal', 'está', 'en', 'la', 'cueva']
```

<!-- grade_id: cell-4350a2848f7fb265, points: 2, grade: true, solution: false -->
```python
assert words == words2
print('Test passed!')
```

Output:
```
Test passed!
```

## Q3 (2 points)

Imagine we want a function that will take a list of words and return a list with each word reversed, like this:

```python
>>> reverse_words(['one', 'two', 'three'])
['eno', 'owt', 'eerht']
```

Below are three ways you could do that (recall that `s[::-1]` is one way to reverse a string `s`):

```python
def reverse_words1(words):
    if words:
        return [words[0][::-1]] + reverse_words1(words[1:])
    else:
        return []

def reverse_words2(words):
    rev_words = []
    for word in words:
        rev_words.append(word[::-1])
    return rev_words

def reverse_words3(words):
    for word in words:
        yield word[::-1]
```

For each of these functions, say which is:

* an accumulating function
* a generator function
* a recursive function

<!-- grade_id: cell-68993764d628fd55, points: 2, grade: true, solution: true -->
reverse_words1 is a recursive function

reverse_words2 is an accumulating function

reverse_words3 is a generator function

## Q4 (1 point)

For the three functions in [Q3](#Q3), note that they appear to return the same results:

```python
>>> list(reverse_words1(['one', 'two', 'three']))
['eno', 'owt', 'eerht']

>>> list(reverse_words2(['one', 'two', 'three']))
['eno', 'owt', 'eerht']

>>> list(reverse_words3(['one', 'two', 'three']))
['eno', 'owt', 'eerht']
```

But this is misleading because the results are cast to a list. Which of the functions does not actually return a list?

<!-- grade_id: cell-2732ca443f15eec4, points: 1, grade: true, solution: true -->
reverse_words3

## Q5 (2 points)

All three functions in [Q3](#Q3) apply an operation to each item on a list. Python's built-in [map()](https://docs.python.org/3/library/functions.html#map) is a [higher-order function](https://en.wikipedia.org/wiki/Higher-order_function) that applies some other function to each item on a list, such as the following which maps the `int()` function to a list of number strings in order to cast them to integers:

```python
>>> list(map(int, ['1', '10', '25']))
[1, 10, 25]
```

It seems like a `map()` is a good match for our purpose. But the way we reversed a string (`s[::-1]`) cannot be passed as the function for `map()`:

```python
>>> list(map([::-1], ['one', 'two', 'three']))
  File "<stdin>", line 1
    list(map([::-1], ['one', 'two', 'three']))
              ^
SyntaxError: invalid syntax
```

What is one way you can make string-reversal into a function that `map()` would accept? You may either answer in words or provide an implementation.

<!-- grade_id: cell-9a6fe8dedca4abcc, points: 2, grade: true, solution: true -->
```python
rightway = ['one', 'two', 'three']
wrongway = list(map(lambda x: "".join(reversed(x)), rightway))
wrongway
```

Output:
```
['eno', 'owt', 'eerht']
```

## Q6 (2 points)

For the following scenarios, say if **unit testing**, **regression testing**, **performance testing**, or **stress testing** would help.

* You have inherited a project written in [Perl](https://en.wikipedia.org/wiki/Perl) but your boss wants you to rewrite it in Python. You want to be sure that the Python version gives the same results as the Perl version. What kind of tests can help here, and why?

<!-- grade_id: cell-df5af7508cc859d1, points: 1, grade: true, solution: true -->
Performance testing and regression testing. Performance testing will test the efficiency of the new code as compared to the original code and regression testing will confirm that any changes have not affected existing features (i.e. that the whole program runs smoothly).

* You have ported the Perl project to Python and you have ensured that the results are correct for small inputs, but you cannot confirm with the full test suite because it takes too long and sometimes uses all the system's memory. You want to find out where the code is inefficient. What kinds of tests can help here, and why?

<!-- grade_id: cell-2e3d54b505845d71, points: 1, grade: true, solution: true -->
Unit testing and performance testing. Unit testing tests a single function or method at a time, and performance testing will test the efficiency of the new code as compared to the original code.